# Interact Exercise 4

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
```

```python
from IPython.html.widgets import interact, interactive, fixed
from IPython.display import display
```

## Line with Gaussian noise

Write a function named `random_line` that creates `x` and `y` data for a line with y direction random noise that has a normal distribution $N(0,\sigma^2)$:

$$
y = m x + b + N(0,\sigma^2)
$$

Be careful about the `sigma=0.0` case.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def random_line(m, b, sigma, size=10):
    """Create a line y = m*x + b + N(0,sigma**2) between x=[-1.0,1.0]
    
    Parameters
    ----------
    m : float
        The slope of the line.
    b : float
        The y-intercept of the line.
    sigma : float
        The standard deviation of the y direction normal distribution noise.
    size : int
        The number of points to create for the line.
    
    Returns
    -------
    x : array of floats
        The array of x values for the line with `size` points.
    y : array of floats
        The array of y values for the lines with `size` points.
    """
    x = np.linspace(-1.0, 1.0, size)
    y = m*x + b + np.random.normal(0,sigma**2)
    return x, y
```

<!-- grade_id: interactex04a, points: 2, grade: true, solution: false -->
```python
m = 0.0; b = 1.0; sigma=0.0; size=3
x, y = random_line(m, b, sigma, size)
assert len(x)==len(y)==size
assert list(x)==[-1.0,0.0,1.0]
assert list(y)==[1.0,1.0,1.0]
sigma = 1.0
m = 0.0; b = 0.0
size = 500
x, y = random_line(m, b, sigma, size)
assert np.allclose(np.mean(y-m*x-b), 0.0, rtol=0.1, atol=0.1)
assert np.allclose(np.std(y-m*x-b), sigma, rtol=0.1, atol=0.1)
```

Output:
```
<ipython-input-39-a2ba28099c0c> in random_line(m, b, sigma, size)
     21     """
     22     x = np.linspace(-1.0, 1.0, size)
---> 23     y = m*x + b + np.random.normal(0,sigma**2)
     24     return x, y

mtrand.pyx in mtrand.RandomState.normal (numpy/random/mtrand/mtrand.c:12477)()

ValueError: scale <= 0
```

Write a function named `plot_random_line` that takes the same arguments as `random_line` and creates a random line using `random_line` and then plots the `x` and `y` points using Matplotlib's `scatter` function:

* Make the marker color settable through a `color` keyword argument with a default of `red`.
* Display the range $x=[-1.1,1.1]$ and $y=[-10.0,10.0]$.
* Customize your plot to make it effective and beautiful.

```python
def ticks_out(ax):
    """Move the ticks to the outside of the box."""
    ax.get_xaxis().set_tick_params(direction='out', width=1, which='both')
    ax.get_yaxis().set_tick_params(direction='out', width=1, which='both')
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def plot_random_line(m, b, sigma, size=10, color='red'):
    """Plot a random line with slope m, intercept b and size points."""
    x = [-1.0, 1.0]
    y = float(m * x + b)
```

```python
plot_random_line(5.0, -1.0, 2.0, 50)
```

Output:
```
<ipython-input-30-b659ae715dab> in <module>()
----> 1 plot_random_line(5.0, -1.0, 2.0, 50)

<ipython-input-29-418d7e6137f5> in plot_random_line(m, b, sigma, size, color)
      2     """Plot a random line with slope m, intercept b and size points."""
      3     x = [-1.0, 1.0]
----> 4     y = float(m * x + b)

TypeError: can't multiply sequence by non-int of type 'float'
```

<!-- grade_id: interactex04b, points: 4, grade: true, solution: false -->
```python
assert True # use this cell to grade the plot_random_line function
```

Use `interact` to explore the `plot_random_line` function using:

* `m`: a float valued slider from `-10.0` to `10.0` with steps of `0.1`.
* `b`: a float valued slider from `-5.0` to `5.0` with steps of `0.1`.
* `sigma`: a float valued slider from `0.0` to `5.0` with steps of `0.01`.
* `size`: an int valued slider from `10` to `100` with steps of `10`.
* `color`: a dropdown with options for `red`, `green` and `blue`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
interact(plot_random_line, m = (-10.0, 10.0, 0.1))
interact(plot_random_line, b = (-5.0, 5.0, 0.1))
interact(plot_random_line, sigma = (0.0, 5.0, 0.01))
interact(plot_random_line, size = (10, 100, 10))
interact(plot_random_line, color = ('red','green','blue'))
```

Output:
```
/usr/local/lib/python3.4/dist-packages/IPython/html/widgets/interaction.py in interactive(__interact_f, **kwargs)
    188     kwargs = kwargs.copy()
    189 
--> 190     new_kwargs = _find_abbreviations(f, kwargs)
    191     # Before we proceed, let's make sure that the user has passed a set of args+kwargs
    192     # that will lead to a valid call of the function. This protects against unspecified

/usr/local/lib/python3.4/dist-packages/IPython/html/widgets/interaction.py in _find_abbreviations(f, kwargs)
    145         for name, value, default in _yield_abbreviations_for_parameter(param, kwargs):
    146             if value is empty:
--> 147                 raise ValueError('cannot find widget or abbreviation for argument: {!r}'.format(name))
    148             new_kwargs.append((name, value, default))
    149     return new_kwargs

ValueError: cannot find widget or abbreviation for argument: 'b'
```

<!-- grade_id: interactex04c, points: 4, grade: true, solution: false -->
```python
assert True # use this cell to grade the plot_random_line interact
```