# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    x = 0
    y = 0
    for i in x:
        for j in y:
            if i == j: #Thought this would check to see if any of the integers in x would be equal to any of those in y.
                return True
            else: 
                return False
            
same_digits(12, 21)
print (same_digits)
```

Output:
```
<ipython-input-18-48545053db19> in <module>()
     10                 return False
     11 
---> 12 same_digits(12, 21)
     13 print (same_digits)

<ipython-input-18-48545053db19> in same_digits(x, y)
      3     x = 0
      4     y = 0
----> 5     for i in x:
      6         for j in y:
      7             if i == j: #Thought this would check to see if any of the integers in x would be equal to any of those in y.

TypeError: 'int' object is not iterable
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-14-4d2369f19b4d> in <module>()
----> 1 assert same_digits('132', '321')
      2 assert not same_digits('123', '3')
      3 assert not same_digits('456', '0987654321')

AssertionError:
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(a, b, c, d, e, f):
    """Do the integers x and y have the same digits, regardless of order."""
    a = 0
    b = 2*a
    c = 3*a
    d = 4*a
    e = 5*a
    f = 6*a
    
    for i in a:
        for j in b:
            for k in c:
                for l in d:
                    for m in e:
                        for n in f:
                            if i == j == k == l == m ==n: #Thought this would check to see if any of the integers in a would be equal to those in b, c, d, e, or f. 
                                return True
                            else: 
                                return False
                            
same_digits()
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-25-791820b7e67d> in <module>()
     19                                 return False
     20 
---> 21 same_digits(5)

TypeError: same_digits() missing 5 required positional arguments: 'b', 'c', 'd', 'e', and 'f'
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```

```python

```