# Numpy Exercise 2

## Imports

```python
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
```

## Factorial

Write a function that computes the factorial of small numbers using `np.arange` and `np.cumprod`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def np_fact(n):
    """Compute n! = n*(n-1)*...*1 using Numpy."""
    x = np.arange(1, n + 1 , 1)
    z = x.cumprod()
    print x
    print z
    
np_fact(10)
```

Output:
```
[ 1  2  3  4  5  6  7  8  9 10]
[      1       2       6      24     120     720    5040   40320  362880
 3628800]
```

<!-- grade_id: numpyex02a, points: 4, grade: true, solution: false -->
```python
assert np_fact(0)==1
assert np_fact(1)==1
assert np_fact(10)==3628800
assert [np_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]
```

Output:
```
[]
[]
AssertionError                            Traceback (most recent call last)
<ipython-input-93-c7ed3614e378> in <module>()
----> 1 assert np_fact(0)==1
      2 assert np_fact(1)==1
      3 assert np_fact(10)==3628800
      4 assert [np_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]

AssertionError:
```

Write a function that computes the factorial of small numbers using a Python loop.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def loop_fact(n):
    """Compute n! using a Python for loop."""
    i = 1
    factorial = 1
    while (i <= n):
        factorial = factorial * i
        i = i + 1

    print factorial
    
loop_fact(10)
```

Output:
```
3628800
```

<!-- grade_id: numpyex02b, points: 4, grade: true, solution: false -->
```python
assert loop_fact(0)==1
assert loop_fact(1)==1
assert loop_fact(10)==3628800
assert [loop_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]
```

Output:
```
1
AssertionError                            Traceback (most recent call last)
<ipython-input-97-1e3be026136e> in <module>()
----> 1 assert loop_fact(0)==1
      2 assert loop_fact(1)==1
      3 assert loop_fact(10)==3628800
      4 assert [loop_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]

AssertionError:
```

Use the `%timeit` magic to time both versions of this function for an argument of `50`. The syntax for `%timeit` is:

```python
%timeit -n1 -r1 function_to_time()
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
import timeit

def np_fact(n):
    """Compute n! = n*(n-1)*...*1 using Numpy."""
    x = np.arange(1, n + 1 , 1)
    z = x.cumprod()
    print z
    
z = np_fact(10)

%timeit -n1 -r1 np_fact(z)
```

Output:
```
[      1       2       6      24     120     720    5040   40320  362880
 3628800]
<magic-timeit> in inner(_it, _timer)

<ipython-input-91-41584a86b212> in np_fact(n)
      3 def np_fact(n):
      4     """Compute n! = n*(n-1)*...*1 using Numpy."""
----> 5     x = np.arange(1, n + 1 , 1)
      6     z = x.cumprod()
      7     print z

TypeError: unsupported operand type(s) for +: 'NoneType' and 'int'
```

```python
def loop_fact(n):
    """Compute n! using a Python for loop."""
    i = 1
    factorial = 1
    while (i <= n):
        factorial = factorial * i
        i = i + 1

    print factorial
    
z = loop_fact(10)

%timeit -n1 -r1 loop_fact(z)
```

Output:
```
3628800
1
1 loops, best of 1: 78 µs per loop
```

In the cell below, summarize your timing tests. Which version is faster? Why do you think that version is faster?

<!-- grade_id: numpyex02c, points: 2, grade: true, solution: true -->
The loop_fact shoule be faster because instead of having to calculate the factorial of every number in the array it only does it for the input value. 