# Project Euler: Problem 17

https://projecteuler.net/problem=17

If the numbers 1 to 5 are written out in words: one, two, three, four, five, then there are 3 + 3 + 5 + 4 + 4 = 19 letters used in total.

If all the numbers from 1 to 1000 (one thousand) inclusive were written out in words, how many letters would be used?


NOTE: Do not count spaces or hyphens. For example, 342 (three hundred and forty-two) contains 23 letters and 115 (one hundred and fifteen) contains 20 letters. The use of "and" when writing out numbers is in compliance with British usage.

First write a `number_to_words(n)` function that takes an integer `n` between 1 and 1000 inclusive and returns a list of words for the number as described above

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
#have to include how to write each number
def number_to_words(n):
    """Given a number n between 1-1000 inclusive return a list of words for the number."""
    for i in n:
        while n<=1000:
            n = 0
            n += 1
            x += len(n)
        print (x)
number_to_words(0)
```

Output:
```
<ipython-input-11-f81510b42311> in <module>()
      7             x += len(n)
      8         print (x)
----> 9 number_to_words(1000)

<ipython-input-11-f81510b42311> in number_to_words(n)
      1 def number_to_words(n):
      2     """Given a number n between 1-1000 inclusive return a list of words for the number."""
----> 3     for i in n:
      4         while n<=1000:
      5             n = 0

TypeError: 'int' object is not iterable
```

```python
ones_place = {0: 0, 1: "one", 2: "two", 3: "three", 4: "four", 5: 'five', 6: '3', 7: 'five', 8: 'five', 9: 'four'}
tens_place = {0: 0, 10: 'ten', 11: 'eleven', 12: 'twelve', 13: 'thirteen', 14: 'fourteen', 15: 'fifteen', 16: 'sixteen', 17: 'seventeen', 18: 'eighteen', 19: 'nineteen', 2: 'twenty', 3: 'thirty', 4: 'forty', 5: 'fifty', 6: 'sixty', 7: 'seventy', 8: 'eighty', 9: 'ninety'}
hundred__place_and = 10

def number_to_words(n):
    if len(n) == 1:
        return ones_place[int(n)]
    elif len(n) == 2:
        if 10 <= int(n) <= 19:
            return tens_place[int(n)]
        else:
            return tens_place[int(n[0])] + ones_place[int(n[1])]
    else:
        x = n[1:2:]
        if int(x) == 1:
            return ones_place[int(n[0])] + hundred__place_and + tens_place[int(n[1::])]
        else:
            if int(n) % 100 != 0:
                return ones_place[int(n[0])] + hundred__place_and + tens_place[int(n[1:2:])] + ones_place[int(n[2])]
            else:
                return ones_place[int(n[0])] + hundred__place_and + tens_place[int(n[1:2:])] + ones_place[int(n[2])] - 3

            
            
if __name__ == '__main__':
    letters = 11
    for n in range(1, 1000):
        letters += convert_to_words(str(n))
    print(letters)
```

Output:
```
21124
```

Now write a set of `assert` tests for your `number_to_words` function that verifies that it is working as expected.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: projecteuler17a, points: 4, grade: true, solution: false -->
```python
assert True # use this for grading the number_to_words tests.
```

Now define a `count_letters(n)` that returns the number of letters used to write out the words for all of the the numbers `1` to `n` inclusive.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
ones_place = {0: 0, 1: "one", 2: "two", 3: "three", 4: "four", 5: 'five', 6: '3', 7: 'five', 8: 'five', 9: 'four'}
tens_place = {0: 0, 10: 'ten', 11: 'eleven', 12: 'twelve', 13: 'thirteen', 14: 'fourteen', 15: 'fifteen', 16: 'sixteen', 17: 'seventeen', 18: 'eighteen', 19: 'nineteen', 2: 'twenty', 3: 'thirty', 4: 'forty', 5: 'fifty', 6: 'sixty', 7: 'seventy', 8: 'eighty', 9: 'ninety'}
hundred__place_and = 10

def count_letters(n):
    if len(n) == 1:
        return ones_place[int(n)]
    elif len(n) == 2:
        if 10 <= int(n) <= 19:
            return tens_place[int(n)]
        else:
            return tens_place[int(n[0])] + ones_place[int(n[1])]
    else:
        x = n[1:2:]
        if int(x) == 1:
            return ones_place[int(n[0])] + hundred__place_and + tens_place[int(n[1::])]
        else:
            if int(n) % 100 != 0:
                return ones_place[int(n[0])] + hundred__place_and + tens_place[int(n[1:2:])] + ones_place[int(n[2])]
            else:
                return ones_place[int(n[0])] + hundred__place_and + tens_place[int(n[1:2:])] + ones_place[int(n[2])] - 3


```

Now write a set of `assert` tests for your `count_letters` function that verifies that it is working as expected.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: projecteuler17b, points: 4, grade: true, solution: false -->
```python
assert True # use this for grading the count_letters tests.
```

Finally used your `count_letters` function to solve the original question.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: projecteuler17c, points: 2, grade: true, solution: false -->
```python
assert True # use this for gradig the answer to the original question.
```